## Übung 3.1

Vergleichen Sie die zahlenmäßig ausgewerteten („quantitativen“) Ergebnisse von 

$$ 
d_0 = a \cdot b + a \cdot c
$$

und 

$$ 
d_1 = a \cdot (b + c)
$$

für $a = 100$, $b = \text{0,1}$ und $c = \text{0,2}$. Speichern Sie $d_{0}$ in der Variablen `d0` und $d_{1}$ in der Variablen `d1`.


Versuchen Sie anschließend, die beiden berechneten Ergebnisse auf Gleichheit zu überprüfen, z. B. `print(d0 == d1)`. 

<!-- grade_id: cell-447110c225124fc4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-91f2d89d88ea86c0, points: 0, grade: true, solution: false -->
```python
assert d0 == 30.0
assert d1 != 30.0
assert d0 != d1
```

## Übung 3.2

Werten Sie mit Python das folgende Polynom aus: 
\begin{align}
f(x, y) &= (x + y)^{6} 
\\
&=  x^6 + 6x^{5}y + 15x^{4}y^{2} + 20x^{3}y^{3} + 15x^{2}y^{4} + 6xy^{5} + y^{6}
\end{align}
(i) in der kompakten Form $(x + y)^{6}$; und (ii) in der ausmultiplizierten Langform für:

(a) $x = 10$ und $y = \text{10,1}$

(b) $x = 10$ und $y = -\text{10,1}$

und vergleichen Sie die Anzahl der signifikanten Stellen, für die für (i) und (ii) die berechneten Ergebnisse übereinstimmen. Speichern Sie das Ergebnis der kompakten Form in einer Variablen mit Bezeichner `f0` und das der ausmultiplizierten Langform in einer Variablen mit Bezeichner `f1`.

Vergleichen Sie für den Variablenbelegungsfall (b) die von Python ausgewerteten Ergebnisse mit Ihrer analytischen Lösung und bestimmen Sie den [relativen Fehler](https://de.wikipedia.org/wiki/Fehlerschranke).
Welche der beiden Rechenwege (i) oder (ii) würden Sie zur numerischen Berechnung des Polynoms mit Python empfehlen?

#### (a) $x = 10$ und $y = \text{10,1}$

<!-- grade_id: cell-e8428e3de87d5500, points: 0, grade: false, solution: true -->
```python
x = 10.0
y = 10.1

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-55d184514edf3a67, points: 0, grade: true, solution: false -->
```python
import math
assert math.isclose(f0, 65944160.60120103, rel_tol=1e-10)
assert math.isclose(f1, 65944160.601201, rel_tol=1e-10)
```

#### (b) $x = 10$ und $y = -\text{10,1}$

<!-- grade_id: cell-1561d20281ff67f6, points: 0, grade: false, solution: true -->
```python
x = 10.0
y = -10.1

# YOUR CODE HERE
raise NotImplementedError()
```

```python
import math
assert math.isclose(f0, 1.0e-6, rel_tol=1e-10)
assert math.isclose(f1, 1.0e-6, rel_tol=1e-2)
```

## Übung 3.3

Gegeben sei der Rechenausdruck

$$
f = \frac{1}{\sqrt{x^2 - 1} - x}
$$

Wenn die Variable $x$ zählenmäßig sehr groß wird, nähert sich der Nenner der Zahl Null, was zu Problemen führen kann.

Versuchen Sie den Rechenausdruck für $f$ so umzuformulieren, dass der Bruch eleminiert wird, indem Sie Zähler und Nenner mit $\sqrt{x^2 - 1} + x$ multiplizieren. Werten Sie die beiden Formulierungen von $f$ aus für:

(a) $x = 1 \cdot 10^{7}$

(b) $x = 1 \cdot 10^{9}$ (Möglicherweise erhalten Sie für diesen Fall einen Python-Fehler.
Warum?)

#### (a) $x = 1 \cdot 10^{7}$

<!-- grade_id: cell-149ac8791e7ff020, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

#### (b) $x = 1 \cdot 10^{9}$

<!-- grade_id: cell-4ef986ef1f5838ba, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```