NumPy zur Verfügung stellen mit:

```python
import numpy as np
```

## Übung 7.1 (Indexierung und Zeitmessung)

Erstellen Sie zwei sehr große NumPy-Felder (auch bezeichnet als „*Arrays*“ oder Vektoren) `x` und `y` mit der gleichen Anzahl an Elementen und summieren Sie diese mit Hilfe:

1. des NumPy-Additionsoperators für Felder, `z = x + y`; und
2. einer `for`-Schleife, mit der Sie die Vektorsumme elementweise berechnen.

Vergleichen Sie die benötigte Zeit der beiden Berechnungstechniken für einige zumehmend umfangreicher definierte Vektoren `x` und `y` (mit zuletzt sehr umfangreichen Vektoren mit beispielsweise zehn Millionen Elementen für die Zeitmessung). Die numerischen Werte der Elemente der beiden an der Summenbildung beteiligten Felder sind für diesen Test bzw. unseren experimentellen Vergleich der beiden Additionstechniken bedeutungslos bzw. nicht relevant.
Nutzen Sie für die Zeitmessung die Python-Funktion `time()` aus dem Modul `time` der Python-Standardbibliothek.

*Hinweis:* Verwenden Sie eine Konstruktion wie die folgende, um indiziert über die Elemente eines Feldes zu iterieren und die Zeitmessung für den zu vergleichenden Anweisungsblock durchzuführen:

```python
from time import time

x = np.ones(10000000)
y = np.ones(len(x))
z = np.zeros(len(x))

t1 = time()
for i in range(len(x)):
    z[i] = x[i] * y[i]
t2 = time()

delta_t = t2 - t1
```

#### (1) Addieren Sie zwei Vektoren gleicher Dimension mit Hilfe des NumPy-Additionsoperators:

<!-- grade_id: cell-2c856e54f7c3340e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

#### (2) Addieren Sie dieselben beiden Vektoren elementweise mit einer `for`-Schleife:

<!-- grade_id: cell-f34614f9f0068cc4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Optionale Erweiterung: Gerade noch rechtzeitige Übersetzung (noch unmittelbar vor der Python-Interpretation und Anweisungsausführung)  eines Python-Code-Abschnittes unmittelbar in die native Maschinensprache des der Programmausführung zugrundeliegenden Prozessors bzw. Digitalrechners (engl. *just-in-time [JIT] compilation*)

Im vorausgegangenen Basisteil zu dieser Übung werden Sie einen großen Unterschied bezüglich der erforderlichen Rechenzeit zwischen Ihrer NumPy- und Ihrer „direkten“ (unmittelbaren) Python-Implementierung der Vektoraddition beobachtet haben. Dies ist darauf zurückzuführen, dass Python eine [*interpretierte*](https://de.wikipedia.org/wiki/Interpreter) im Gegensatz zu einer [*kompilierten*](https://de.wikipedia.org/wiki/Compiler) (übersetzten, d. h. einer bereits im Voraus noch vor der Programmausführung in Maschinensprache überführten) Programmiersprache ist. Eine Möglichkeit, unmittelbare Python-Implementierungen zu beschleunigen, besteht darin, den interpretierten Python-Code in kompilierten Maschinen-Code umzuwandeln. Ein Software-Werkzeug dafür ist [Numba](https://numba.pydata.org/).

Unten sehen Sie ein Beispiel für die Verwendung von Numba und JIT zur Beschleunigung einer Berechnung:

```python
import numba
import math

def berechne_Sinus_unmittelbar_in_Python(x):
    z = np.zeros(len(x))
    for i in range(len(z)):
        z[i] = math.sin(x[i])
    return z

@numba.jit
def berechne_Sinus_mit_JIT_Uebersetzung(x):
    z = np.zeros(len(x))
    for i in range(len(z)):
        z[i] = math.sin(x[i])
    return z
    
x = np.ones(10000000)
t1 = time()
z = berechne_Sinus_unmittelbar_in_Python(x)
t2 = time()
z = berechne_Sinus_mit_JIT_Uebersetzung(x)
t3 = time()
print("Delta_t interpretiert =", t2 - t1)
print("Delta_t mit Numba JIT =", t3 - t2)
```

**Aufgabe:** Testen Sie, ob und um wie viel Numba zur Beschleunigung Ihrer obigen unmittelbaren Implementierung der Vektoraddition mittels elementweiser Indizierung verwendet werden kann.

## Übung 7.2 (Mitgliedsfunktionen [„*member functions*“] von NumPy-Feldern und Extraktion einer Teilansicht (einer flachen Kopie bzw. „Scheibe“, engl. *slice*, einer zusammenhängenden, geordneten Teilmenge von Feldelemementen als „Unterfeld[ansicht]“) mittels „*slicing*“)

Prüfungsergebnisse werden für jeden anonymisierten Prüfling als erreichte Punktzahl (von insgesamt sechzig erzielbaren Gesamtpunkten) in einem NumPy-Array gespeichert. Implementieren Sie:

1. eine Python-Funktion, die ein NumPy-Feld der erreichten Punktzahlen als formalen Parameter erwartet und ein NumPy-Feld der erreichten Prozentwerte in aufsteigend sortierter Reihenfolge  zurückgibt. (Testen Sie dafür die NumPy-Array-Mitgliedsfunktion `ergebnisse.sort()`, wobei `ergebnisse` ein NumPy-Array ist, das die erreichten Punktzahlen der Prüflinge enthält.)
1. eine Python-Funktion, die das Minimum, Maximum und den arithmetischen Mittelwert der erreichten Punktzahlen berechnet und diese Maßzahlen (auch bezeichnet als beschreibende bzw. deskriptive Statistik) als ein Nachschlagewerk vom Python-Datentyp `dict` mit den Schlüsseln `min`, `max` und `mean` zurückgibt. Verwenden Sie die NumPy-Array-Mitgliedsfunktionen `min()`, `max()` und `mean()` für die Berechnung der Maßzahlen, z. B.  
`max = ergebnisse.max()`.  
   
   Entwerfen Sie Ihre Python-Funktion für die Berechnung der beschreibenden Statistik der erreichten Punktzahlen so, dass Sie optional die höchste und die niedrigste erreichte Punktzahl (als statistische „Ausreißer“) von der Berechnung von Minimum, Maximum und arithmetischem Mittelwert ausschließen können. 
   
   *Hinweis:* sort the array of scores and use array slicing to exclude
   the first and the last entries.

Verwenden Sie für den Test Ihrer Python-Funktion die folgenden erreichten Punktzahlen:
```python
ergebnisse = np.array([58.0, 35.0, 24.0, 42, 7.8])
```

<!-- grade_id: cell-169ebae60810c6be, points: 0, grade: false, solution: true -->
```python
def berechne_aufsteigend_sortierte_prozentwerte(ergebnisse):
    # YOUR CODE HERE 
    raise NotImplementedError()

def berechne_beschreibende_statistik(ergebnisse, ausschluss=False):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-af0b6fd8a3cadb1a, points: 0, grade: true, solution: false -->
```python
ergebnisse = np.array([58.0, 35.0, 24.0, 42, 7.8])
assert np.isclose(berechne_aufsteigend_sortierte_prozentwerte(ergebnisse), [ 13.0, 40.0, 58.33333333,  70.0, 96.66666667]).all()

s0 = berechne_beschreibende_statistik(ergebnisse)
assert round(s0["min"] - 7.8, 10) == 0.0
assert round(s0["mean"] - 33.36, 10) == 0.0
assert round(s0["max"] - 58.0, 10) == 0.0

s1 = berechne_beschreibende_statistik(ergebnisse, True)
assert round(s1["min"] - 24.0, 10) == 0.0
assert round(s1["mean"] - 33.666666666666666667, 10) == 0.0
assert round(s1["max"] - 42.0, 10) == 0.0
```

## Übung 7.3 (slicing)

For the two-dimensional array

```python
A = np.array([[4.0, 7.0, -2.43, 67.1],
             [-4.0, 64.0, 54.7, -3.33],
             [2.43, 23.2, 3.64, 4.11],
             [1.2, 2.5, -113.2, 323.22]])
print(A)
```

use array slicing for the below operations, printing the results to the screen to check. Try to use array slicing such that your code would still work if the dimensions of `A` were enlarged.



#### 1. Extract the third column as a 1D array

<!-- grade_id: cell-dfee6b0ed9343682, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

#### 2. Extract the first two rows as a 2D sub-array

<!-- grade_id: cell-7bf2f9a8c67029f8, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

#### 3.  Extract the bottom-right $2 \times 2$ block as a 2D sub-array

<!-- grade_id: cell-5206fea47d246222, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

#### 4. Sum the last column

<!-- grade_id: cell-34f74988e47b9f87, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

#### Compute transpose

Compute the transpose of `A` (search online to find the function/syntax to do this).

<!-- grade_id: cell-dd2999d2da8070f6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Übung 7.4 (optional extension)

In a previous exercise you implemented the bisection algorithm to find approximate roots of a mathematical function. Use the SciPy bisection function `optimize.bisect` (http://docs.scipy.org/doc/scipy-0.14.0/reference/generated/scipy.optimize.bisect.html) to find roots of the mathematical function that was used in the previous exercise. Compare the results computed by SciPy and your program from the earlier exercise, and compare the computational time (using `%time`).

<!-- grade_id: cell-b6580accbcb3c2da, points: 0, grade: false, solution: true -->
```python
from scipy import optimize

# YOUR CODE HERE
raise NotImplementedError()
```