## Übung 6.1 (Auswahl einer geeigneten Datenstruktur zur vereinfachten Parameterübergabe an eine Funktion)

In Übung 4.3 haben Sie eine Python-Funktion entwickelt, die den Flächeninhalt eines Dreickes aus sechs formalen Parametern (den $x$- und $y$-Koordinaten der drei Eckpunkte des Dreickes) berechnet. Bei sechs formalen Parametern ist die Wahrscheinlichkeit groß, dass ein Aufrufer der Funktion beim Funktionsaufruf die aktuellen Parameter (die „Argumente“) in der falschen Reihenfolge übergibt. 

Verwenden Sie eine geeignete Datenstruktur, wie z. B. `list`, `tuple` oder `dict`, um eine neue Version Ihrer Funktion mit einer vereinfachten Schnittstelle zu entwickeln. (Die Schnittstelle definiert die formalen Parameter einer Funktion.) Ergänzen Sie in der Definition Ihrer neuen Funktion auch geeignete Tests, um die Gültigkeit der an die vereinfachten formalen Parameter übergebenen Eingabedaten zu überprüfen.

<!-- grade_id: cell-50c4566d7bcbffaf, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Übung 6.2 (Auswahl geeigneter Datenstrukturen)

Der Flächeninhalt $A$ eines einfachen Polygons („einfach“ ist ein Polygon, wenn sich seine Kanten nicht schneiden) mit den $n$ aufeinanderfolgenden Eckpunkten gegeben durch deren Koordinaten $(x_0, y_0)$, $(x_1, y_1)$, . . , $(x_{n-1}, y_{n-1})$ lässt sich mit der folgenden Formel berechnen:
$$
A = \left| \frac{1}{2} \sum_{i=0}^{n-1} \left(x_{i} y _{i+1} - x_{i+1} y_{i} \right) \right|
$$
wobei $(x_n, y_n) = (x_0, y_0)$. Die Eckpunkte des Polygons müssen dabei beim Umlauf entlang der Kanten des Polygons in einer gewählten Richtung in genau ihrer auftretenden Reihenfolge durchnummeriert werden. Der gewählte erste Eckpunkt (als Startpunkt) erhält dabei den Index $0$ und entspricht nach vollständigem Umlauf auch gleichzeitig dem letzten Eckpunkt (dem Zielpunkt) mit Index $n$.

Entwickeln Sie eine Python-Funktion zur Berechnung des Flächeninhaltes eines einfachen Polygons mit beliebig vielen Eckpunkten. Testen Sie die Ihre Funktion für einige einfache Polygone. Achten Sie bei der Konstruktion Ihrer Testfälle darauf, dass Sie keine Polygone mit Schleifen oder sich schneidenen Kanten entwerfen. 

<!-- grade_id: cell-4b5a619405787746, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Übung 6.3 (Indexierung)

Schreiben Sie eine Python-Funktion, die die Indexierung von Listenelementen verwendet, um zwei übergebene Vektoren beliebiger Länge addieren zu können und einen neuen Vektor (den Summenvektor) an ihren Aufrufer zurückgibt. Fügen Sie Ihrer Funktionsdefinition eine Kontrolle hinzu, um zu überprüfen, dass die beiden übergebenen Vektoren nicht leer sind und die gleiche Anzahl an Elementen haben und geben Sie andernfalls einen entsprechenden Warnhinweis (bzw. eine „Fehlermeldung“) auf der Konsole aus. Je spezifischer Ihre Fehlermeldung ist, desto einfacher wird es für einen Aufrufer Ihrer Funktion, die Fehlersuche im aufrufenden Python-Code durchzuführen.

Fügen Sie Ihrer Lösung bzw. Implementierung einige geeignete Testfälle hinzu.

#### Implemtierungshinweis: Sie können eine Liste aus `n` Nullen erstellen mit

    z = [0] * n
    
#### Optional (fortgeschritten bzw. weiterführend, für besonders Interessierte) 

Versuchen Sie eine alternative, einzeilige Version ihrer zuvor entwickelten Vektoraddition mittels der sog. „*Listen-Abstraktion*“ (engl. *list comprehension*) und der eingebauten Python-Funktion [`zip`](https://docs.python.org/3/library/functions.html#zip) zu erstellen.

<!-- grade_id: cell-47daca03f7a96b62, points: 0, grade: false, solution: true -->
```python
def addiere_vektoren(x, y):
    """"Bildet den (elementweisen) Summenvektor aus (den als Python-Listen) übergebenen
        Vektoren x und y und gibt diesen (als Python-Liste) an den Aufrufer zurück."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-95f5819810d2b2c7, points: 0, grade: true, solution: false -->
```python
a = [0, 4.3, -5, 7]
b = [-2, 7, -15, 1]
c = addiere_vektoren(a, b)
assert c == [-2, 11.3, -20, 8]
```

### Optionale Erweiterung: „*Listen-Abstraktion*“ (engl. *list comprehension*)

<!-- grade_id: cell-c182b3f0d8952692, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Übung 6.4 (Datenstruktur für „Wörterbücher“, engl. *dictionaries*)

Erstellen Sie mit Hilfe der Datenstruktur „Wörterbuch“, die in Python durch den eingebauten Datentyp `dict` bereitgestellt wird, ein Verzeichnis bzw. ein „Nachschlagewerk“ (Lexikon), das *Namen von deutschen Bundesländern* (als sog. „Schlüssel“, engl. *key*) auf deren offizielle Abkürzung (hier das „Nachschlageobjekt“ bzw. den „-wert“, engl. *value*) für mindestens fünf Bundesländer abbildet
(Sie finden diese Abkürzungen z. B. beim [Statistischen Bundesamt](https://www.destatis.de/DE/Themen/Laender-Regionen/Regionales/Gemeindeverzeichnis/Glossar/bundeslaender.html) oder international auch als [ISO-Standard](https://de.wikipedia.org/wiki/ISO_3166-2:DE)).
Aus Ihrem Python-Wörterbuch erstellen und geben Sie auf der Konsole aus:

1. ein weiteres Verzeichnis, das Bundeslandkürzel (als „*key*“) auf Bundeslandnamen (als „*value*“) abbildet; und
1. eine Liste der gespeicherten Bundeslandabkürzungen sortiert in alphabetisch aufsteigender Reihenfolge.

*Optionale Erweiterung:* Erstellen Sie ein „Verzeichnis“ bzw. ein „Nachschlagewerk“ (Lexikon) vom Python-Datentyp `dict`, das jeweils einen Bundeslandnamen (als „Schlüssel“) abbildet auf ein Dreitupel (auch bezeichnet als „Tripel“, hier das „Nachschlageobjekt“ bzw. der „-wert“):

- seines offiziellen Kürzels (seiner „Abkürzungsbezeichung“),
- seiner zuletzt offiziell statisch geschätzten Einwohneranzahl (in Mio.); und
- seiner Fläche in Quadratkilometern (km$^2$)
 
für mindestens fünf Bundesländer. Entnehmen Sie die Daten zu den Bundesländern z. B. der freien Enzyklopädie [Wikipedia](https://de.wikipedia.org/wiki/Land_&#040;Deutschland&#041;). Bestimmen Sie mit Hilfe Ihres Verzeichnisses:
 
1. das Bundesland mit der größten Einwohnerzahl und geben Sie dessen Kürzel auf der Konsole aus; und 
1. das flächenmäßig größte Bundesland und geben Sie dessen Einwohneranzahl (in Mio.) und seine Abkürzungsbezeichnung auf der Konsole aus.

<!-- grade_id: cell-bc271dc3885e19c5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

#### Optionale Erweiterung:

<!-- grade_id: cell-7040c4053c72ce09, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```