## Übung 5.1 (Zufallszahlen, engl. *random numbers*)

- Entwickeln Sie eine Python-Funktion `werfe_wuerfel_mit_seiten`, die den Wurf eines Würfels mit $n$ Seiten [virtuell](https://de.wikipedia.org/wiki/Virtualität) simuliert (d. h. rein im Digitalrechner, was in der Bioinformatik auch lateinisch angelehnt als „[*in silico*](https://de.wikipedia.org/wiki/In_silico)“ bezeichnet wird). Verwenden Sie in Ihrer Implementierung die Funktion `randint` aus dem Modul `random` der Python-Standardbibliothek (https://docs.python.org/3/library/random.html#random.randint), um ganzzahlige sog. [Pseudozufallszahlen](https://de.wikipedia.org/wiki/Pseudozufall) zu generieren.

- Konstruieren (d. h. konzipieren und implementieren) Sie für einen virtuellen Würfel mit $n=6$ Seiten einen Test zur Qualitätssicherung, um zu prüfen, ob es sich bei Ihrer Implementierung um einen *fairen Würfel* handelt.

#### (a) Ihre Implementierung der Funktion `werfe_wuerfel_mit_seiten`:

<!-- grade_id: cell-70a83466894475ee, points: 0, grade: false, solution: true -->
```python
import random

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0970743b0c035f78, points: 0, grade: true, solution: false -->
```python
for n in range(1, 20):
    for j in range(100):
        gewuerfelte_zahl = werfe_wuerfel_mit_seiten(n) 
        assert gewuerfelte_zahl >= 1 and gewuerfelte_zahl <= n
```

#### (b) Ihr Test auf Fairness:

<!-- grade_id: cell-0826760c62c21bff, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Übung 5.2 (Datenkompression)

Die Kompression von Daten kann für Digitalrechner mit begrenztem Speicher oder begrenzter [Datenübertragungsrate](https://de.wikipedia.org/wiki/Datenübertragungsrate) wichtig sein. [Datenkompression](https://de.wikipedia.org/wiki/Datenkompression) ist ein eigenes  Fachgebiet der Informatik, aber mit Funktionsbibliotheken können wir Daten einfach kompremieren („*einpacken*“) und bei Bedarf wieder verlustlos (engl. *lossless*) dekompremieren („*entpacken*“), ohne dass wir selbst Detailwissen aus dieser Fachwissenschaft benötigen.

Der folgende Python-Code kompremiert einen Textauszug aus „*Hamlet - Prinz von Dänemark*“ von William Shakespeare.

```python
# Importiere das Modul zur Datenkompression
# aus der Python-Standardbibliothek
import zlib

# Lege die Zeichenkette an, die wir kompremieren wollen
text = """
Willkommen, Rosenkranz und Güldenstern!
Wir wünschten nicht nur sehnlich, Euch zu sehn,
Auch das Bedürfnis Eurer Dienste trieb
Uns zu der eilgen Sendung an. Ihr hörtet
Von der Verwandlung Hamlets schon; so nenn ichs,
Weil nicht der äußre noch der innre Mensch
Dem gleicht, was sonst er war. Was es nur ist,
Mehr als des Vaters Tod, das ihn so weit
Von dem Verständnis seiner selbst gebracht,
Kann ich nicht raten. Ich ersuch Euch beide,
Da Ihr von Kindheit auf mit ihm erzogen
Und seiner Laun und Jugend nahe bliebt,
Ihr wollet hier an unserm Hof verweilen
Auf einge Zeit, um ihn durch Euren Umgang
In Lustbarkeit zu ziehn und zu erspähn,
Soweit der Anlaß auf die Spur Euch bringt,
Ob irgendwas, uns unbekannt, ihn drückt,
Das, offenbart, zu heilen wir vermöchten."""

# Wandle die (16-Bit Unicode) Python-Zeichenkette
# vom Python-Datentyp 'string' in eine Byte-Folge
# nach dem UTF-8 Standard um und prüfe den Typ
text_unkompremiert = text.encode("utf-8")
print(type(text_unkompremiert))

# Ermittle die Anzahl der Byte, die zum Speichern der Zeichenkette verwendet werden
print("Anzahl Byte der unkompremierten Zeichenkette:", len(text_unkompremiert))

# Kompremiere die Zeichenkette und prüfe die Anzahl der Byte der kompremierten Zeichenkette
text_kompremiert = zlib.compress(text_unkompremiert)
print("Anzahl Byte der kompremierten Zeichenkette  :", len(text_kompremiert))

# Gib die Kompressionsrate auf der Konsole aus
print("Kompressionsrate:", len(text_kompremiert) / len(text_unkompremiert))

# Dekompremiere ("entpacke") die kompremierte Zeichenkette wieder
text_dekompremiert = zlib.decompress(text_kompremiert)

# Prüfe, ob die ursprüngliche und die dekompremierte Zeichenkette inhaltlich übereinstimmen
if text != text_dekompremiert.decode("utf-8"):
    print("Problem: Die ursprüngliche und dekompremierte Zeichenkette unterscheiden sich.")
```

Untersuchen Sie mit dem vorstehenden Python-Code als Orientierungshilfe die Kompressionsrate für

1. die Kompression einer langen Zeichenkette, die sich aus einer einhundertmaligen Aneinanderreihung (Wiederholung) des obigen Hamlet-Textauszuges zusammensetzt; und
2. die Kompression einer Zufallszeichenkette (engl. *random string*), die aus zufällig ausgewählten Zeichen aus einem gegebenen Alphabet von Unicode-Zeichen gebildet wird und die dieselbe Länge (d. h. die gleiche Anzahl an Zeichen) wie die wiederholte (einhundertmal aneinandergereihte) Hamlet-Zeichenkette besitzt.

Als Hilfestellung für Ihre Untersuchung, erzeugt die im folgenden gegebene Python-Funktion eine Zufallszeichenkette bestehend aus `anzahl_zufallszeichen` vielen Zufallszeichen:

```python
import random

def erzeuge_zufallszeichenkette(anzahl_zufallszeichen):
    alphabet = ' abcdefghijklmnopqrstuvwxyzABCDEFGHIJKLMNOPQRSTUVWXYZäöüßÄÖÜ.,;?!'
    zufallszeichenkette = ''
    for i in range(anzahl_zufallszeichen):
        zufallszeichen = random.choice(alphabet) 
        zufallszeichenkette += zufallszeichen
    return zufallszeichenkette

print(erzeuge_zufallszeichenkette(8))
```

### Hilfestellung zur Lösung

```python
# Erstelle die Zeichenkette
text = """
Willkommen, Rosenkranz und Güldenstern!
Wir wünschten nicht nur sehnlich, Euch zu sehn,
Auch das Bedürfnis Eurer Dienste trieb
Uns zu der eilgen Sendung an. Ihr hörtet
Von der Verwandlung Hamlets schon; so nenn ichs,
Weil nicht der äußre noch der innre Mensch
Dem gleicht, was sonst er war. Was es nur ist,
Mehr als des Vaters Tod, das ihn so weit
Von dem Verständnis seiner selbst gebracht,
Kann ich nicht raten. Ich ersuch Euch beide,
Da Ihr von Kindheit auf mit ihm erzogen
Und seiner Laun und Jugend nahe bliebt,
Ihr wollet hier an unserm Hof verweilen
Auf einge Zeit, um ihn durch Euren Umgang
In Lustbarkeit zu ziehn und zu erspähn,
Soweit der Anlaß auf die Spur Euch bringt,
Ob irgendwas, uns unbekannt, ihn drückt,
Das, offenbart, zu heilen wir vermöchten."""
```

Importieren Sie die beiden benötigten Module („Pakete“) aus der Python-Standardbibliothek:

```python
import random
import zlib
```

Reihen Sie den Hamlet-Textauszug einhundertmal aneinander, um eine sehr lange Zeichenkette zu bilden, und kompremieren Sie diese:

<!-- grade_id: cell-c08bb020d9be6620, points: 0, grade: false, solution: true -->
```python
# Erstelle eine sehr lange Zeichenkette, die aus einer einhundertmaligen
# Aneinanderreihung des Hamlet-Textauszuges besteht und kompremieren Sie diese

# YOUR CODE HERE
raise NotImplementedError()
```

Erstellen Sie eine Zufallszeichenkete der gleichen Länge und kompremieren Sie diese:

<!-- grade_id: cell-c5d8c07589ec3cd1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Vergleichen Sie die Kompressionsraten der beiden Zeichenkettenkompremierungen miteinander:

<!-- grade_id: cell-a2dd98ee370e22a3, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```