## Übung 2.1 (if-else-Anweisung)

Gegeben sei das folgende Bewertungsschema, nach welchem die in einer Prüfungsleistung vom Prüfling erreichte Punktzahl auf einen Bewertungskommentar abgebildet wird:

| Bewertungskommentar         | erreichte Punktzahl $p$ (von 100)  |
| --------------------------- | ---------------------------------- |
| hervorragend                | $p \ge 82$                         |
| sehr gut                    | $76,5 \le p < 82$                  |
| gut                         | $66 \le p < 76,5$                  |
| Da ist noch Luft nach oben! | $45 \le p < 66$                    |
| Haben Sie es denn versucht? | $p < 45$                           |

Schreiben Sie ein Python-Programm, das zu einer erreichten Punktzahl den zugehörigen Bewertungskommentar ausgibt.

<!-- grade_id: cell-71121f470be22e4a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Übung 2.2 (Bisektionsverfahren)

Beim Bisektions- bzw. Intervallhalbierungsverfahren handelt es sich um eine Technik, um beispielsweise iterativ nach den Nullstellen einer gegebenen Funktion zu suchen (siehe https://de.wikipedia.org/wiki/Bisektion für einen Überblick über die Anwendungen dieser Technik anhand von Beispielen). Nehmen wir an, die Funktion $f(x)$ hat eine Nullstelle zwischen $x_{0}$ und $x_{1}$ ($x_{0} < x_{1}$). Wir ...

- werten $f$ an der mittigen Zwischenstelle $x_{\rm m} = (x_0 + x_1)/2$ aus, d. h. wir berechnen
   $f_{\rm m} = f(x_{\rm m})$
- berechnen $f(x_0) \cdot f(x_{\rm m})$

  - wenn $f(x_0) \cdot f(x_{\rm m}) < 0$, 

    muss $f$ in seinem Werteverlauf irgendwo zwischen $x_0$ and $x_{\rm m}$ einen Vorzeichenwechsel aufweisen, daher muss eine Nullstelle von $f$ zwischen $x_0$ und $x_{\rm m}$ liegen, weshalb wir $x_{\rm m}$ an $x_1$ zuweisen.
   
  - sonst:

    muss $f$ in seinem Werteverlauf irgendwo zwischen $x_{\rm m}$ und $x_1$ einen Vorzeichenwechsel aufweisen, weshalb wir  $x_{\rm m}$ an $x_0$ zuweisen.

Die obigen Schritte können eine gewisse Anzahl mal oder solange wiederholt ausgeführt werden, bis $|f_{\rm m}|$ unter einer spefizierten (d. h. „ingenieurmäßig“ vorgegebenen) Toleranzschwelle liegt (manchmal technisch auch als sog. „Schwellenwert“ bezeichnet), woraufhin das resultierende $x_{\rm m}$ als angenäherte („approximierte“) Nullstelle ausgegeben wird.


### Aufgabe

Die Funktion

$$
f(x) = x^3 - 6x^2 + 4x + 12
$$

hat eine Nullstelle irgendwo zwischen $x_0 = 3$ und $x_1 = 6$.

1. Wenden Sie das Biseketionsverfahren an, um eine approximierte Nullstelle $x_{n}$ mit 15 Iterationen unter Verwendung einer `for`-Schleife zu suchen.
2. Wenden Sie das Biseketionsverfahren an, um eine approximierte Nullstelle $x_{n}$ unter Verwendung einer `while`-Schleife zu suchen, so dass $\left| f(x_{n}) \right| < 10^{-6}$ ist und geben Sie auch die dafür benötigte Anzahl an Iterationen mit aus.

Speichern Sie die angenäherte Nullstelle in der Variablen `x_m` und speichern Sie $f(x_{\rm m})$ in einer Variablen mit dem Bezeichner `f`.

*Hinweis:* Verwenden Sie `abs` um den „Absolutbetrag“ (kurz: den „Betrag“) einer Zahl zu berechnen, z. B. `y = abs(x)` weist den Absolutbetrag des numerischen Inhaltes der Variablen mit Bezeichner `x` an die Variable mit Bezeichner `y` zu. 

#### (1) Verwendung einer `for`-Schleife:

<!-- grade_id: cell-3a0d0a59613d6b3f, points: 0, grade: false, solution: true -->
```python
# Initiale (d. h. "anfängliche") Intervallgrenzen
x0 = 3.0
x1 = 6.0

# Verwende 15 Wiederholungen ("Iterationen")
for n in range(15):
    # Berechne die Intervallmitte (d. h. die mittige Zwischenstelle zwischen den Intervallgrenzen)
    x_m = (x0 + x1) / 2

    # Werte die Funktion an der linken Intervallgrenze (x0) und in der Intervallmitte (x_m) aus
    f0 = x0**3 - 6*x0**2 + 4*x0 + 12
    f = x_m**3 - 6*x_m**2 + 4*x_m + 12
    
    # YOUR CODE HERE
    raise NotImplementedError()

    print(n, x_m, f)
```

<!-- grade_id: cell-421824bbd4e61fae, points: 0, grade: true, solution: false -->
```python
assert round(x_m - 4.534149169921875, 10) == 0.0
assert abs(f) < 0.0009
```

#### (2) Verwendung einer `while`-Schleife:

Nutzen Sie die Variable `zaehler`, um die Anzahl der Iterationen mitzuzählen. 

*Denken Sie auch an den Mechanismus zum Schutz vor einer Endlosschleife.*

<!-- grade_id: cell-1747a3fb1dd1de48, points: 0, grade: false, solution: true -->
```python
# Initiale Intervallgrenzen
x0 = 3.0
x1 = 6.0

tolerenz = 1.0e-6
abweichung = toleranz + 1.0

# Iterate until tolerance is met
zaehler = 0
while abweichung > toleranz:
    # YOUR CODE HERE
    raise NotImplementedError()

    # Schutzmechanismus vor einer Endlosschleife
    if zaehler > 1000:
        print("Uups, der Iterationszähler hat einen sehr großen Wert. Ausbruch aus der while-Schleife.")
        break
    
    print(zaehler, x_m, abweichung)
```

<!-- grade_id: cell-ddff32f44df88388, points: 0, grade: true, solution: false -->
```python
assert zaehler == 23
assert abs(f) < 1.0e-6
```

## Übung 2.3 (Reihenentwicklung)

Eine Potenzreihenentwicklung der Sinusfunktion (zum Begriff der „Potenzreihe“ siehe z. B. https://de.wikipedia.org/wiki/Reihenentwicklung als Ausgangspunkt für die Hintergründe) ist gegeben durch die folgende Formel: 

$$
\sin(x) = \sum_{n = 0}^{\infty} (-1)^n \frac{x^{2n +1}}{(2n+1)!}
$$

1. Unter Verwendung einer `for`-Schleife: Nähern Sie $\sin(\frac{3\pi}{2})$ mit den ersten 15 Summanden der oben gegebenen Potenzreihe für $ \sin(x)$ an und berichten Sie den verbleibenden „absoluten Fehler“ (d. h. den Absolutbetrag des Fehlers).

1. Unter Verwendung einer `while`-Schleife: Berechnen Sie wie viele Summanden der oben gegebenen Potenzreihenformel für $\sin(x)$ erforderlich sind, um $\sin(\frac{3\pi}{2})$ unterhalb einer Toleranzgrenze von $10^{-8}$ für den absoluten Fehler zu approximieren. 

Speichern Sie den Absolutbetrag des Fehlers in der Variablen `fehler`.

*Hinweis:* Taschenrechner und auch die leistungsfähigeren „Computer“ (alles Arten von sog. „Digitalrechnern“ als informatischer Oberbegriff) verwenden im Allgemeinen derartige Iterationsverfahren oder Reihenentwicklungen, um beispielsweise die Werte trigonometrischer Funktionen zu berechnen, ganz ähnlich wie in den obigen beiden Aufgabenstellungen, wenngleich sie auch in der Regel effizientere Formulierungen als die im obigen Beispiel angegebene Potenzreihenformel für $\sin(x)$ verwenden.

### Hinweis

Verwenden Sie das Modul `math` aus der Python-Standardbibliothek, um die Fakultät zu berechnen und eine gute Näherung für die der Kreiszahl $\pi$ zu erhalten:
```python
import math
n = 10
n_fakultaet = math.factorial(n)
pi = math.pi
```
Sie müssen die Anweisung `import math` nur einmal zu Beginn Ihres Programms geben. Modulewie  `math`aus der Python-Standardbibliothek werden wir später noch behandeln. Um den Wert der Sinusfunktion für nicht einfache Winkel im Bogenmaß zu bestimmen, können Sie sich am folgenden Beispiel orientieren:
```python
alpha = 1.3
sinus_alpha = math.sin(alpha)
```    
um damit den absoluten Fehler Ihrer eigenen Implentierung gegenüber einer genauen Berechnung der Sinusfunktion zu ermitteln.

#### (1)  Verwendung einer `for`-Schleife:

<!-- grade_id: cell-d6df4b428f969a5d, points: 0, grade: false, solution: true -->
```python
# Importiere das Standardmodul zur Mathematik von Python für den Zugriff auf math.sin and math.factorial
import math

# Winkelwert im Bogenmaß, für den die Sinusunktion ausgewertet werden soll
x = 1.5 * math.pi

# Initiale Approximation der Sinusfunktion
approx_sin = 0.0

# YOUR CODE HERE
raise NotImplementedError()
    
print("Der absolute Fehler ist:")
print(fehler)
```

<!-- grade_id: cell-8b269b717e92414d, points: 0, grade: true, solution: false -->
```python
assert fehler < 1.0e-12
```

#### (2) Verwendung einer `while`-Schleife:

*Denken Sie auch an den Schutz vor einer Endlosschleife.*

<!-- grade_id: cell-03df25d83df3674e, points: 0, grade: false, solution: true -->
```python
# Importiere das Standardmodul zur Mathematik von Python für den Zugriff auf math.sin and math.factorial
import math

# Winkelwert im Bogenmaß, für den die Sinusunktion ausgewertet werden soll
x = 1.5 * math.pi

# Toleranz und initialer absoluter Fehler, der nur größer als die Toleranz sein muss
toleranz = 1.0e-8
fehler = toleranz + 1.0

# Initiale Approximation der Sinusfunktion
approx_sin = 0.0

# Anfangswert des Zählers für die Wiederholungen
n = 0

# Loop until error satisfies tolerance, with a check to avoid 
# an infinite loop
while fehler > toleranz and n < 1000:
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    # Increment counter
    n += 1    
    
    
print("Der absolute Fehler ist:")
print(fehler)

print("Die Anzahl erforderlicher Summanden der Reihenentwicklung von sin(x) war:")
print(n)
```

<!-- grade_id: cell-03df25d83df3674dfwe, points: 0, grade: true, solution: false -->
```python
assert fehler <= 1.0e-8
```