## Übung 1.1

Grad Fahrenheit ($T_F$) lassen sich mit der folgenden Formel in Grad Celsius ($T_C$) umrechnen:

$
T_C = 5(T_F - 32)/9
$

Schreiben Sie ein Programm, das 78 Grad Fahrenheit in Grad Celsius umrechnet und das Ergebnis in der Konsole ausgibt.
Schreiben Sie Ihr Programm so, dass es einfach ist, den umzurechnenden Grad Fahrenheit Wert zu verändern, ohne den Code für die Umrechnungsformel selbst ändern zu müssen.

Verwenden Sie eine Variable mit dem Namen `T_C` für den Temperaturwert in Grad Celsius.

<!-- grade_id: cell-ca2d6af7fabf2918, points: 0, grade: false, solution: true -->
```python
T_F = 78
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a7542313c8cdf4ed, points: 0, grade: true, solution: false -->
```python
assert round(T_C - 25.55555555555555, 10) == 0.0
```

## Übung 1.2

Sie wurden mit der Entwicklung eines Software-Werkzeuges zur Prüfung der Tragfähigkeit von beantragten Krediten beauftragt. Eine Komponente davon besteht darin, die monatlichen Zinszahlungen für eine Reihe von Szenarien zu berechnen.

Die Zinsen für einen Kredit werden zu einem festen Satz über dem „offiziellen Zinssatz“ der Europäischen Zentralbank (EZB) erhoben. Die Zinsen werden jährlich berechnet, und die Zinszahlungen werden gleichmäßig über jeden Monat des Jahres verteilt. Schreiben Sie ein Programm, das die monatlich zu zahlenden Zinsen berechnet mit Variablen für:

1. Darlehenskapital (geliehener Betrag),
1. Offizieller („nominaler“) (Leit-)Zinssatz der EZB (Prozentsatz, ausgedrückt pro Jahr, lat. „per annum“, abgekürzt „p. a.“),
1. Erhöhungswert (als Prozentsatz, ausgedrückt pro Jahr) über den Leitzinssatz der EZB, der zusätzlich durch das Kreditinstitut zur Deckung der Kreditkosten für das Kreditinstitut (z. B. von Bearbeitungsgebühren oder Zinsfestschreibungsdauern) veranschlagt wird,
1. Laufzeit des Kredites (in Jahren).

Testen Sie Ihr Programm mit einem Darlehensbetrag von 150.000 €, einem offiziellen EZB-Leitzinssatz von 0,5 % und einem Erhöhungswert von 1,49 % über den offiziellen EZB-Leitzinssatz durch das Kreditinsitut. Verwenden Sie den Variablennamen `monatlicher_zinsbetrag` für den monatlich zu zahlenden Zinsbetrag.

<!-- grade_id: cell-aaab53918bb9627d, points: 0, grade: false, solution: true -->
```python
# Zuweisung von konkreten Kreditparametern an Variablen
darlehenskapital = 150000 # in Euro

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-258e3a2c80f6a765, points: 0, grade: true, solution: false -->
```python
assert round(monatlicher_zinsbetrag - 248.75, 10) == 0.0
```